```python
# Initialize Otter
import otter
grader = otter.Notebook("lab2.ipynb")
```

![](img/563_lab_banner.png)

# Lab 2: Dimensionality Reduction

## Imports <a name="im"></a>

```python
import pickle
from hashlib import sha1

%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from PIL import Image
from sklearn import datasets
from sklearn.decomposition import NMF, PCA
from sklearn.preprocessing import StandardScaler
```

<br><br><br><br>

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions <a name="si"></a>
rubric={mechanics}

You will receive marks for correctly submitting this assignment by following the instructions below:
    
- Be sure to follow the [general lab instructions](https://ubc-mds.github.io/resources_pages/general_lab_instructions/).
- [Here](https://github.com/UBC-MDS/public/tree/master/rubric) you will find the description of each rubric used in MDS.
- Make at least three commits in your lab's GitHub repository.    
- Push the final .ipynb file with your solutions to your GitHub repository for this lab.        
- Before submitting your lab, run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Clear All Outputs` and then `Run -> Run All Cells`. Notebooks without the output displayed may not be graded at all (because we need to see the output in order to grade your work).     
- Make sure to enroll to Gradescope via [Canvas](https://canvas.ubc.ca/courses/106525).
- Upload the .ipynb file to Gradescope.
- Make sure that your plots/output are rendered properly in Gradescope.    
- If the .ipynb file is too big or doesn't render on Gradescope for some reason, also upload a pdf (preferably WebPDF) or html export of .ipynb file with your solutions so that TAs can view your submission on Gradescope. 
- The data you download for this lab <b>SHOULD NOT BE PUSHED TO YOUR REPOSITORY</b> (there is also a `.gitignore` in the repo to prevent this).
- Include a clickable link to your GitHub repo for the lab just below this cell.
</div>    

_Points:_ 3

YOUR REPO LINK GOES HERE

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 1: Warm-up 
<hr>

### 1.1 Dimensionality reduction notation 
rubric={accuracy}

**Your tasks:**

Suppose you apply dimensionality reduction on a data matrix $X_{n \times d}$ using a dimensionality reduction technique such as PCA with $k$ component, where 

- $n \rightarrow $ number of examples 
- $d \rightarrow $ number of features
- $k \rightarrow $ number of components         

Based on the Principal Component Analysis (PCA) notation we saw in class, what would be the dimensions of each of the matrices `W`, `Z`, and `X_hat`? 
    
- (A) $n \times k$ 
- (B) $n \times d$
- (C) $k \times d$

<div class="alert alert-warning">

Solution_1_1
    
</div>

_Points:_ 2

```python
# Assign the variables to either "A", "B", or "C" 
W_dim = ...
Z_dim = ...
X_hat_dim = ...
```

```python
grader.check("q1.1")
```

<br><br>

<!-- BEGIN QUESTION -->

### 1.2 Picking the number of components $k$
rubric={reasoning}

**Your tasks:**
    
If you want to reduce dimensionality, why is it not a good idea to pick the number of components $k$ (`n_components` in `sklearn`) which gives the lowest reconstruction error?     

<div class="alert alert-warning">

Solution_1_2
    
</div>


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 1.3 PCA by hand

Suppose you train the standard PCA algorithm on an already centered data matrix `X` (not shown) and you get `Z` and `W` shown below.    

```python
Z = np.array([[10, 10], [5, 2], [4, 3], [4, 3]])
W = np.array([[0.5, 0.5, -0.5, -0.5], [0.7, 0.1, 0.7, 0.1]])
```

```python
Z
```

```python
W
```

#### 1.3.1 
rubric={accuracy}

**Your tasks:**

1. How many rows and columns would be there in the data matrix `X`? 

<div class="alert alert-warning">

Solution_1_3_1
    
</div>

_Points:_ 1

```python
X_rows = ...
X_cols = ...
```

```python
grader.check("q1.3.1")
```

<br><br>

#### 1.3.2 
rubric={accuracy}

**Your tasks:**
    
Fill in the blanks: 

- In this toy example, we are reducing dimensionality from **$d$** dimensions to **$p$** dimensions. 


<div class="alert alert-warning">

Solution_1_3_2
    
</div>

_Points:_ 1

```python
# original dimensions
d = ...

# reduced_dimensions
p = ...
```

```python
grader.check("q1.3.2")
```

<br><br>

#### 1.3.3
rubric={accuracy}

Find the low-dimensional representation of the already centered `X_new` below.    

```python
X_new = np.array([[1,1,1,1],[1,0,1,1]])
```

<div class="alert alert-warning">

Solution_1_3_3
    
</div>

_Points:_ 1

```python
Z_new = ...
```

```python
grader.check("q1.3.3")
```

<br><br>

<!-- BEGIN QUESTION -->

#### 1.3.4
rubric={reasoning}
    
The third and fourth rows of the transformed matrix `Z` and the original matrix `X` are identical in our toy example. Does this mean that for two arbitrary data points in a dataset if the lower dimensional representations are identical, the higher dimensional representations in the original space have to be identical? Briefly explain.  

```python
Z
```

<div class="alert alert-warning">

Solution_1_3_4
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 1.4 Reconstruction error 

Let's get an intuition for reconstruction error on a toy dataset.

The code below creates a toy dataset with a few outliers. The function `get_recon_error_df` calculates normalized reconstruction errors between the original and reconstructed data points. Run the code and answer the following questions. 

```python
np.random.seed(42)
outliers = np.array([[4, -3], [2.8, -3], [-3, 3]])
x1 = np.random.randn(100)
x2 = x1 + np.random.randn(100) / 2
X = np.stack([x1, x2]).T
X_noise = np.vstack([X, outliers])
X_noise_scaled = StandardScaler().fit_transform(X_noise)
plt.scatter(X_noise_scaled[:, 0], X_noise_scaled[:, 1], edgecolors='black', c='xkcd:azure');
```

<!-- BEGIN QUESTION -->

#### 1.4.1
rubric={reasoning}

**Your tasks:**    

1. Write a docstring for the function `get_recon_error_df` provided below.

<div class="alert alert-warning">

Solution_1_4_1

</div>


_Points:_ 2

```python
def get_recon_error_df(orig, recon):
    """
    ...
    """
    loss = np.sqrt(np.sum((orig - recon) ** 2, axis=1))
    loss = (loss - np.min(loss)) / (np.max(loss) - np.min(loss))  # normalization
    loss_df = pd.DataFrame(data=loss, columns=["recon_error"])
    return loss_df
```

<!-- END QUESTION -->

<br><br>

#### 1.4.2
rubric={accuracy}

**Your tasks:**    

1. Fit `PCA` with `n_components=1` on `X_noise_scaled` with `random_state=123` and store the returned dataframe in `recon_df` below. 

<div class="alert alert-warning">

Solution_1_4_2
    
</div>

_Points:_ 2

```python
...
```

```python
recon_df = ...
recon_df
```

```python
grader.check("q1.4.2")
```

<br><br>

#### 1.4.3
rubric={accuracy}

**Your tasks:**    

1. The last 3 rows (rows from indices 100 to 102) in `X_noise_scaled` are outliers. Calculate the average reconstruction error for outliers vs. average reconstruction error for other data points (i.e., non-outliers). 

<div class="alert alert-warning">

Solution_1_4_3
    
</div>

_Points:_ 1

_Type your answer here, replacing this text._

```python
outliers_avg_recon_error = ...
non_outliers_avg_recon_error = ...
```

```python
...
```

```python
grader.check("q1.4.3")
```

<br><br><br><br>

## Exercise 2: Implementing PCA using SVD
<hr>
rubric={accuracy}

In this exercise, you'll implement your own version of PCA using SVD. The class `MyPCA` below implements `init` and `fit` methods. 

**Your tasks:** 
    
1. Complete the `get_components` method of the class which returns the learned components. 
2. Complete the `transform` method of the class which returns transformed `Z` given `X`. 
> *Before applying transformation, center the data by subtracting the mean.*
3. Complete `reconstruct` method of the class which returns reconstructed `X_hat` given transformed `Z`.
> *Do not forget to add the mean back after reconstruction.*
4. Run your code and compare results of your PCA and `sklearn` PCA using the code below.  
   

<div class="alert alert-warning">

Solution_2
    
</div>

_Points:_ 8

_Type your answer here, replacing this text._

```python
class MyPCA:
    """
    Solves the PCA problem min_Z,W (Z*W-X)^2 using SVD
    """

    def __init__(self, k):
        self.k = k
        

    def fit(self, X):
        self.mean = np.mean(X, 0)
        X = X - self.mean  # Centralize the data
        U, S, Vt = np.linalg.svd(
            X
        )  # SVD to get singular values and principal components
        self.W = Vt[: self.k, :]  # store only first k components in self.W

        
    def get_components(self):
        """
        Returns principal components.

        Parameters
        -----------
        None

        Returns
        -----------
        np.ndarray: an array containing k principal components.
        """
        ### Solution_2_1
        ...

        
    def transform(self, X):
        """
        Transforms X to Z and return Z.

        Parameters
        -----------
        X : np.ndarray
            Data to be transformed

        Returns
        -----------
        np.ndarray: transformed data
        """
        ### Solution_2_2

        ...

        
    def reconstruct(self, Z):
        """
        Given transformed data Z, returns reconstructed X_hat.

        Parameters
        -----------
        Z : np.ndarray
            PCA transformed data

        Returns
        -----------
        np.ndarray: X_hat which has dimensions of original X
        """
        ### Solution_2_3

        ...
```

```python
from sklearn.datasets import make_blobs

X, y = make_blobs(n_samples=100, centers=3, n_features=20)  ## Generating toy data

for i in range(1, X.shape[1] + 1):
    print("PCA implementation with {} components: OK".format((i)))
    pca = PCA(n_components=i)
    pca.fit(X)

    mypca = MyPCA(k=i)
    mypca.fit(X)

    assert np.allclose(
        np.abs(pca.components_), np.abs(mypca.get_components())
    ), "W values do not match"

    Z = pca.transform(X)
    Z_prime = mypca.transform(X)

    assert np.allclose(np.abs(Z), np.abs(Z_prime)), "Z values do not match"

    X_hat = pca.inverse_transform(Z)
    X_hat_prime = mypca.reconstruct(Z_prime)
    assert np.allclose(
        np.abs(X_hat), np.abs(X_hat_prime)
    ), "reconstructed X_hat values do not match"
```

```python
grader.check("q2")
```

<br><br><br><br>

## Dimensionality reduction on animal faces 
<hr>

In the next few exercises, you'll explore dimensionality reduction on a subset of [Animal Faces dataset](https://www.kaggle.com/andrewmvd/animal-faces) from Kaggle. I have created a subset of this dataset and preprocessed it a bit. 
    
**Your tasks:**    
    
- Download animal_faces.pkl from [here](https://github.ubc.ca/mds-2021-22/datasets/blob/master/data/animal_faces.pkl) and save it locally under the data folder in this lab's directory. 
- Run the code below to unpickle the data and to display the first few images.   

```python
import pickle

animals = pickle.load(open("data/animal_faces.pkl", "rb"))
```

```python
import matplotlib as mpl

mpl.rcParams.update(mpl.rcParamsDefault)
plt.rcParams["image.cmap"] = "gray"
```

```python
animals.shape
```

```python
fig, axes = plt.subplots(2, 5, figsize=(12, 5), subplot_kw={"xticks": (), "yticks": ()})
for image, ax in zip(animals, axes.ravel()):
    ax.imshow(image)
plt.show()
```

Let's flatten the images. 

```python
X_anims = animals.reshape(len(animals), -1)
X_anims.shape
```

The flattened representation of the data is of shape $1500 \times 10000$; Each image is represented with 10,000 features, each feature representing a pixel from the image. Let's define `image_shape` variable which will be handy later when we want to display images. Use this flattened representation `X_anims` in the rest of the lab. 

```python
image_shape = (100, 100)
```

<br><br>

## Exercise 3: Dimensionality reduction with PCA
<hr>

In this exercise, you will explore dimensionality reduction on the animal faces dataset using PCA. 

<!-- BEGIN QUESTION -->

### 3.1 Choosing the number of components 
rubric={accuracy,reasoning}   

First, let's pick the appropriate value for the number of components. Recall that PCA finds principal components such that the first principal component has the highest variance, the second component has the next highest variance and so on. One way to decide the value of number of principal components ($k$ or `n_components`) is based on the amount of variance explained with $k$ components. 

**Your tasks:**

1. Using [scikit-learn's `PCA`](https://scikit-learn.org/stable/modules/generated/sklearn.decomposition.PCA.html) implementation and the following hyperparameter values, fit a PCA model, and plot $k$ (for $k=1, \ldots, 300$) vs. the proportion of total variance explained by the first $k$ components. 
    - `n_components=300`
    - `random_state=42`
2. What range of values for $k$ (`n_components` in `scikit-learn`) seems reasonable based on this plot? Briefly explain.         

<div class="alert alert-info">

_Note that scikit-learn's PCA does the data centering for you. **Do not scale the data in this exercise and the next exercises**._ 

</div>

<div class="alert alert-warning">

Solution_3_1
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.2 Applying PCA
rubric={accuracy}   

**Your tasks:**

1. Train a pca model with the number of components you chose in 3.1 and `random_state=42`.    
2. Get $Z$ (the transformed data), $W$ (principal components), and $X_{hat}$ (reconstructions).     

<div class="alert alert-warning">

Solution_3_2
    
</div>

_Points:_ 4

```python
...
```

```python
Z = ...
W = ...
X_hat = ...
```

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 4: Interpreting the components learned by PCA
<hr>

PCA is a linear transformation and we can visualize and interpret the principal components learned by PCA. In this exercise, you will attempt to manually inspect and interpret these components. 

<!-- BEGIN QUESTION -->

### 4.1 Visualizing PCA components
rubric={viz,reasoning}

Principal component matrix returned by PCA ($W$ matrix) is of shape: number of components by number of features. So you can reshape each component and visualize it as an image (in our case $100 \times 100$ image). 

**Your tasks:**
1. Visualize the first 15 components as $100 \times 100$ images in a grid of 3 rows and 5 columns.
2. Observe the components and identify some possible semantic themes captured by at least four components of your choice. 

<div class="alert alert-info">
    
_Feel free to use code from lecture notes with appropriate attributions._
    
</div>    

<div class="alert alert-warning">

Solution_4_1
    
</div>

_Points:_ 6

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 4.2 Plotting strong component images
rubric={accuracy,reasoning}

**Your tasks:**

1. Pick four principal components from 4.1 where you observed some clear semantic themes. For these principal components, extract and display a few example images where the component values are strong for these particular components. 
2. Briefly comment on your results. Do the example images agree with your interpretation from 4.1?  

<div class="alert alert-info">
    
Hint: The function to extract and show strong component images is provided below.  
    
</div>


```python
def plot_strong_comp_images(X_anims, Z, W, compn=1, image_shape=(100,100)):
    """
    Given transformed data Z, components W, and the component index compn, the function 
    extracts and shows the images from X_anims where the component compn is strongest.   

    Parameters
    ----
    X_anims : numpy array 
        Input data
        
    Z : numpy array
        Transformed data
    
    W : numpy array
        Components found by dimensionality reduction

    compn : int
        component to examine

    image_shape : tuple
        shape of the input image
    Returns
    ----
    None. Shows the images with strong value for the given component
    """    
    inds = np.argsort(Z[:, compn])[::-1]    
    fig, ax = plt.subplots(
        1, 7, figsize=(12,4), subplot_kw={"xticks": (), "yticks": ()}
    )    
    ax[0].set_title(f"Component {compn}")
    ax[0].imshow(W[compn].reshape(image_shape))
    i = 1
    for image in inds[:6]:
        ax[i].set_title(image)
        ax[i].imshow(X_anims[image].reshape(image_shape))
        i+=1
    fig.tight_layout()  
    plt.show()
```

<div class="alert alert-warning">

Solution_4_2
    
</div>

_Points:_ 6

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 4.3 Visualization of first two dimentions
rubric={viz}    

One of the use cases of PCA is visualization. It's not possible to visualize high dimensional data. That said, since the first couple of components of PCA usually capture a lot of information from the data, we can plot the first two dimensions to get an intuition about the patterns in the data. 
    
**Your tasks:**       

1. Make a scatterplot of the first two dimensions in the transformed data $Z$ (from 3.2).       

<div class="alert alert-warning">
    
Solution_4_3
    
</div>

_Points:_ 2

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 4.4 Image tiling 
rubric={accuracy,quality,reasoning}    

In this exercise you will attempt to interpret the first two dimensions of the transformed data. One way to interpret these dimensions is as follows:
- Create an $m \times m$ grid which roughly spans scatterplot region from the previous exercise. For example, if `Z` is your transformed data, the following range of values will get you five points spanning the first dimension. 
```
np.linspace(np.min(Z[:, 0]), np.max(Z[:, 0]), 5))
```
- Once you have representative points which span the first dimension, get the indices of the data points closest to these points. 
- Plot the images corresponding to these indices as a grid and observe whether you see any pattern as the first dimension increases (left to right of the grid) and as the second dimension increases (bottom to top of the grid). 

Let's try this out with $m=5$, i.e., a $5 \times 5$ grid. 
    
**Your tasks:**
    
1. Make a $5 \times 5$ grid that roughly spans the scatterplot region from the previous exercise (but stays within it). For each point on that grid, select the animal face whose first two dimensions are closest to the grid point. Plot a $5 \times 5$ tiling of these animal faces, corresponding to the $5 \times 5$ grid using the `img_tiling` function below.     
2. What happens to the images as the first dimension increases (i.e., go from left to right of the grid)? What about the second dimension? Briefly discuss your observations and try to label your axes based on your interpretation.

<div class="alert alert-info">
    
Hint: The function to make the $5 \times 5$ tiling plot is provided below   
    
</div>


```python
def img_tiling(idx, size=10, image_shape=(100, 100)):
    """
    Plots a 5x5 tiling of faces from bottom to top and left to right. 
    
    Parameters:
    -----------
    idx: the indexes of the faces to be plotted. This should be a 5x5 matrix, where each
         elements is an index corresponding to the closest animal face of that grid point.
         position 0,0 of idx corresponds to the bottom left position in the grid 
         position 0,4 of idx corresponds to the top left position
         position 4,0 of idx corresponds to the bottom right position
         position 4,4 of idx corresponds to the top right position

    size: the desired size of the plot;
    """
    idx = np.array(idx, dtype="int32")  # Just making sure the indexes are int

    plt.figure(figsize=(size, size))  # Creating the image with the desired size

    tile_size = 5
    # Ploting the 5x5 tiling
    for i in range(tile_size):
        for j in range(tile_size):
            face = np.reshape(
                animals[idx[i, j]], (image_shape)
            )  # Obtain the closest face
            plt.imshow(
                face, extent=(i * 32, (i + 1) * 32, j * 32, (j + 1) * 32)
            )  # Plot the closest animal face
    plt.xlim((0, 160))
    plt.ylim((0, 160))
    plt.xticks([])
    plt.yticks([])
```

<div class="alert alert-warning">
    
Solution_4_3
    
</div>

_Points:_ 10

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 5: Non-negative Matrix Factorization (NMF)
<hr>

<!-- BEGIN QUESTION -->

### 5.1 Dimensionality reduction with NMF
rubric={accuracy}

**Your tasks:**

1. Carry out dimensionality reduction with [NMF](https://scikit-learn.org/stable/modules/generated/sklearn.decomposition.NMF.html) with `n_components=15` and `random_state=42` and create transformed data `Z_nmf`. 

<div class="alert alert-warning">
    
Solution_5_1
    
</div>

_Points:_ 2

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 5.2 NMF components 
rubric={accuracy}

**Your tasks:**

1. Write code to display NMF components and show images corresponding to some of the interesting components using the function `plot_strong_comp_images` from 4.2.  

<div class="alert alert-warning">
    
Solution_5_2
    
</div>

_Points:_ 4

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 5.3 Interpretation of NMF components
rubric={reasoning}

**Your tasks:**

1. Interpret at least 4 components of your choice and the corresponding images. In what scenarios would NMF be preferred over PCA? 

<div class="alert alert-warning">
    
Solution_5_3
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 6: Food for thought
<hr>

Each lab will have a few challenging questions. These are usually low-risk questions and will contribute to maximum 5% of the lab grade. The main purpose here is to challenge yourself, dig deeper in a particular area, and going beyond what we explicitly discussed in the class. When you start working on labs, attempt all other questions before moving to these challenging questions. If you are running out of time, please skip the challenging questions. 

![](img/eva-game-on.png)

<br><br>

<!-- BEGIN QUESTION -->

### (Challenging) 6.1 Reconstruction error for anomaly detection
rubric={reasoning}

**Your tasks:**

1. Write a paragraph on how PCA and reconstruction errors might be used for anomaly detection. 
2. Get reconstruction errors for images using the function `get_recon_error_df` from Exercise 1. Write code to display a few original and reconstructed image pairs, where the reconstruction error is very high and where it is very low.
3. Are you able to find anomalous images based on reconstruction error. Briefly discuss. 

<div class="alert alert-info">
    
You might want to look up robust PCA for additional information on this topic.    
    
</div>    

<div class="alert alert-warning">

Solution_6.1
    
</div>

_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### (Challenging) 6.2 Clustering animal faces   
<hr>

rubric={reasoning}

**Your tasks:**

1. Reduce dimensionality of the data using PCA and explore clustering animal faces with KMeans or other clustering methods of your choice with this representation.  
2. Comment on the clustering results. 

<div class="alert alert-warning">
    
Solution_6_2
    
</div>

_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### (Challenging) 6.3 Latent Semantic Analysis (LSA)
<hr>

rubric={reasoning}

**Your tasks:**
- Extract topics from the recipe titles dataset we used in lab1 using LSA ([TruncatedSVD](https://scikit-learn.org/stable/modules/generated/sklearn.decomposition.TruncatedSVD.html)). Show the topics and discuss your results. 

<div class="alert alert-warning">
    
Solution_6_3
    
</div>

_Points:_ 1

_Type your answer here, replacing this text._

```python
...
```

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br><br><br>

Before submitting your assignment, please make sure you have followed all the instructions in the Submission Instructions section at the top. 

Well done!! Have a wonderful weekend! 

```python
from IPython.display import Image

Image("img/eva-happy-caturday.png")
```

Tests for q1.1 (points: 2)
```
>>> assert sha1(W_dim.encode('utf8')).hexdigest() == '32096c2e0eff33d844ee6d675407ace18289357d', "W_dim is incorrect"
>>> assert sha1(Z_dim.encode('utf8')).hexdigest() == '6dcd4ce23d88e2ee9568ba546c007c63d9131c1b', "Z_dim is incorrect"
>>> assert sha1(X_hat_dim.encode('utf8')).hexdigest() == 'ae4f281df5a5d0ff3cad6371f76d5c29b6d953ec', "X_hat_dim is incorrect"
```


Tests for q1.3.1 (points: 1)
```
>>> assert sha1(str(X_rows).encode('utf8')).hexdigest() == '1b6453892473a467d07372d45eb05abc2031647a', "Incorrect number of rows."
>>> assert sha1(str(X_cols).encode('utf8')).hexdigest() == '1b6453892473a467d07372d45eb05abc2031647a', "Incorrect number of columns."
```


Tests for q1.3.2 (points: 1)
```
>>> assert sha1(str(d).encode('utf8')).hexdigest() == '1b6453892473a467d07372d45eb05abc2031647a', "Your answer is incorrect."
```
failure_message: Incorrect original dimensions
```
>>> assert sha1(str(p).encode('utf8')).hexdigest() == 'da4b9237bacccdf19c0760cab7aec4a8359010b0', "Your answer is incorrect."
```
failure_message: Incorrect reduced dimension


Tests for q1.3.3 (points: 1)
```
>>> assert sha1(str(Z_new.shape).encode('utf8')).hexdigest() == 'a0975468caa87203bf270355e1a6c54c2f724db7', "Your answer is incorrect."
```
failure_message: Z_new shape is incorrect.
```
>>> assert sha1(str(Z_new[0][0]).encode('utf8')).hexdigest() == '38f6d7875e3195bdaee448d2cb6917f3ae4994af', "Your answer is incorrect."
>>> assert sha1(str(Z_new[0][1]).encode('utf8')).hexdigest() == '4693695d02f6236c431a6589af6dfdef6aa7fa31', "Your answer is incorrect."
```
failure_message: Incorrect elements in Z_new


Tests for q1.4.2 (points: 2)
```
>>> assert sha1(str(recon_df.shape[0]).encode('utf8')).hexdigest() == '934385f53d1bd0c1b8493e44d0dfd4c8e88a04bb', "Your answer is incorrect."
```
failure_message: Incorrect number of rows
```
>>> assert sha1(str(np.round(recon_df.iloc[0][0], 3)).encode('utf8')).hexdigest() == 'f48b8680bbd951f42a2f865d51e3e8eeaec7142b', "Your answer is incorrect."
```
failure_message: Incorrect number of cols


Tests for q1.4.3 (points: 1)
```
>>> assert sha1(str(np.round(outliers_avg_recon_error, 3)).encode('utf8')).hexdigest() == 'dec89deff3a3ecde32b11029fcf20550a8b6fec2', "Your answer is incorrect."
```
failure_message: Average of reconstruction error of outliers is incorrect.
```
>>> assert sha1(str(np.round(non_outliers_avg_recon_error, 3)).encode('utf8')).hexdigest() == 'ff30105317e5aa2bcbbe271e4617e7e61d9c6859', "Your answer is incorrect."
```
failure_message: Average of reconstruction error of non outliers is incorrect.


Tests for q2 (points: 8)
```
>>> 
>>> def test_it():
...     X, y = make_blobs(n_samples=200, centers=4, n_features=30, random_state=123)
...     for i in range(1, X.shape[1] + 1):
...         pca = PCA(n_components=i)
...         pca.fit(X)
...         mypca = MyPCA(k=i)
...         mypca.fit(X)
...         assert np.allclose(np.abs(pca.components_), np.abs(mypca.get_components())), "W values do not match"
...         Z = pca.transform(X)
...         Z_prime = mypca.transform(X)
...         assert np.allclose(np.abs(Z), np.abs(Z_prime)), "Z values do not match"
...         X_hat = pca.inverse_transform(Z)
...         X_hat_prime = mypca.reconstruct(Z_prime)
...         assert np.allclose(np.abs(X_hat), np.abs(X_hat_prime)), "reconstructed X_hat values do not match"
...         
>>> test_it()
```
success_message: Good job!
